---
title: Frees DataCamp Audition Video
description: This is a template chapter.
---

## Ex 1.1

```yaml
type: NormalExercise
lang: r
xp: 100
skills: 1
key: 210d9bb371
```

Do some data science.

`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```

`@sample_code`

```{r}

```

`@solution`

```{r}

```

`@sct`

```{r}

```

---
## Method of Least Squares

```yaml
type: VideoExercise
key: 7928777a1a
lang: r
xp: 50
skills: 1
video_link: player.vimeo.com/video/154783078
```


